# Useful modules in the standard library

<!-- (Adapted from [OOP in Python](https://python-textbok.readthedocs.io/en/1.0/Useful_Libraries.html)) -->

Python comes with many built-in modules, such as the `math` module you encountered in Lab 1. This notebook will introduce you to a few more. There are homework problems throughout. You may solve these in the programming language of your choice and submit a different file, or if using Python, you can do them directly in this notebook.

(Random tip: you can quickly comment/uncomment code with Ctrl+L.)

For a complete list of modules in the standard library, check out [Python Standard Library documentation](https://docs.python.org/3/library/index.html); note that the latest version of Python may have different choices than what you are running. Also note that if you use the Anaconda distribution of Python, there are many more modules pre-installed beyond those in the standard library. You can see what is installed on your computer by opening the Anaconda Prompt (a special command prompt/terminal for Anaconda) and running `conda list`.

> Terminology note: a *package* is what you install to get new *modules*. The programs `pip` and `conda` are *package managers* that can install packages and upgrade existing modules. A *library* is a collection of modules.

[datetime](#datetime)
[math](#math)
[random](#random)
[re](#re)
[csv](#csv)
[Questions](#Questions)

<a id='datetime'></a>

## Date and time: `datetime`

The datetime module provides us with objects which we can use to store information about dates and times:

* `datetime.date` is used to create dates which are not associated with a time.
* `datetime.time` is used for times which are independent of a date.
* `datetime.datetime` is used for objects which have both a date and a time.
* `datetime.timedelta` objects store differences between dates or datetimes – if we subtract one datetime from another, the result will be a timedelta.
* `datetime.timezone` objects represent time zone adjustments as offsets from UTC. This class is a subclass of `datetime.tzinfo`, which is not meant to be used directly.

We can query these objects for a particular component (like the year, month, hour or minute), perform arithmetic on them, and extract printable string versions from them if we need to display them. Here are a few examples:

> Run the code cells to see the results. Feel free to edit cells and re-run to experiment with the code. You can create new cells to run `help` commands on objects for more information.

```python
import datetime

# this class method creates a datetime object with the current date and time
now = datetime.datetime.today()

print(now.year)
print(now.hour)
print(now.minute)
```

Output:
```
2021
20
46
```

```python
print(now.weekday())

print(now.strftime("%a, %d %B %Y"))

long_ago = datetime.datetime(1999, 3, 14, 15, 9, 26)

print(long_ago) # remember that this calls str automatically
print(long_ago < now) # nice that comparison operators are defined on dates
```

Output:
```
1
Tue, 06 July 2021
1999-03-14 15:09:26
True
```

> The full list of options for `strftime` is available [here](https://docs.python.org/3.6/library/datetime.html#strftime-and-strptime-behavior). Note that capital and lowercase options have different meanings!

```python
difference = now - long_ago
print(type(difference))
print(difference) # remember that this calls str automatically
```

(Note the `type` function in the 2nd to last line. That returns what data type a variable is. This can be useful for validating user input, or for exploring variables when you're learning new code.)

```python
user_input = "hello"
if type(user_input) != float:
    print('error')
```

----------

There's a saying in programming:

> *The sooner you start writing code, the longer it takes.*

This means, if you immediately start to write code without spending time planning your work, you'll often spend extra time re-working your code. Try taking a few moments to plan out the code you'll write in plain English for the following exercises. Maybe write some comments to help structure things, then fill in the code afterward.

---------

## <font color="blue">Exercise 1</font>

Print ten dates, each two weeks apart, starting from today. Format them in the form *YYYY-MM-DD*. So, dates only, no times. Each date should appear on its own line by itself.

Hint: use a [loop](https://www.pythonlikeyoumeanit.com/Module2_EssentialsOfPython/ForLoops.html#For-Loops) to make things better. Use the [`range`](https://www.pythonlikeyoumeanit.com/Module2_EssentialsOfPython/Generators_and_Comprehensions.html#The-range-generator) generator.

Since this is homework, please make your answers "stand alone". As in, import whatever modules you need for *each* answer. In practice, we don't do that (just import one time), but for assignments, please do. (Assume I'll run only the cells for an individual exercise to test it.)

Also keep readability in mind...use "good" variable names, write useful comments as appropriate, and such.

Make sure you test your code a bit before submitting. You can use the tests I provide below or write some of your own.

<!-- grade_id: printDates, points: 0, grade: false, solution: true -->
```python
import datetime 

def printDates(date):
    """ Prints the next ten days from a given starting date
        args: Year, month, day
        
    """
# Two strings for clarity of variables set within the function after running the code
    print('Starting Date: {d}'.format(d=date))
    print("Next 10 dates :")
    #Loop to iterate through 10 dates sequentially  
    for _ in range(10):
        date=date+datetime.timedelta(days=14)
        print('{d}'.format(d=date))

dt = datetime.date(2016,8,1)
printDates(dt)
    
```

Output:
```
Starting Date: 2016-08-01
Next 10 dates :
2016-08-15
2016-08-29
2016-09-12
2016-09-26
2016-10-10
2016-10-24
2016-11-07
2016-11-21
2016-12-05
2016-12-19
```

These are automated tests that I will run on your solution. You should run them yourself before submitting by running this cell. Note that I will often run additional tests as appropriate, so your solution needs to solve the problem, not merely pass the provided tests.

<!-- grade_id: printDates_tests, points: 2, grade: true, solution: false -->
```python
#
# This is a function I wrote to help me with grading. You are free to ignore it.
#
def _test_printedFunction(func):
    """Returns a string of any printed output and a single/tuple of any objects returned by func."""
    import io, sys
    
    capturedOutput = io.StringIO()                  # Create StringIO object to store output from print()
    current_stdout = sys.stdout                     # Save current output to reset later
    sys.stdout = capturedOutput                     #  and redirect stdout.
    
    func_res = func()                               # Call function. Printed output will be captured, return values goto func_res
    sys.stdout = current_stdout                     # Reset redirect immediately.
    
    return capturedOutput.getvalue(), func_res      # Return contents of captured ouptut and any func return values

##############################


printDates()  # This just runs the function without checking anything. Actually running is an important first step!
    
# _test_printedFunction() is a custom function I use to help with grading
output, res = _test_printedFunction(printDates)
output = output.split()  # convert into a list; by default, split() operates on line endings


# assert statements: If the condition is True, nothing happens; if False, an exception is raised. 
# You want nothing to happen.
assert len(output) == 10  # function actually printed 10 items
assert len(output[0]) == 10  # the item is of the form YYYY-MM-DD
assert output[0].startswith('2021-07-')  # day left out to let us run this successfully on different days

```

<a id='math'></a>
## Mathematical functions: `math`

You saw this one before, so this is just a reminder. The functions in `math` can be used on floats and integers (mostly floats), and often return floats. Here are a few examples:

```python
import math

# These are constant attributes, not functions
math.pi
math.e

# round a float up or down
math.ceil(3.3)
math.floor(3.3)

# natural logarithm
math.log(5)

# logarithm with base 10
math.log(5, 10)
math.log10(5) # this function is slightly more accurate

# square root
math.sqrt(10)

# trigonometric functions
math.sin(math.pi/2)
math.cos(0)

# convert between radians and degrees
math.degrees(math.pi/2)
math.radians(90)
```

Output:
```
1.5707963267948966
```

(Note that in a Jupyter notebook code cell, the last command is usually printed by default. So if that's all you want, you don't need `print()` statements.)

If you need to do math on complex numbers, use the `cmath` module instead:

```python
import cmath

cmath.log(3.3j + 27)
```

Output:
```
(3.303250763250131+0.12161902326134658j)
```

## <font color="blue">Exercise 2</font>

Write an object that represents a sphere of a given radius. Write a function that calculates the sphere's volume and another function that calculates its surface area.

Spheres with radius=0 are confusing, so don't allow that. In fact, let's force all spheres to have a minimum radius of 1. You can warn the user you're changing the radius, but you shouldn't stop execution.

(Since we haven't covered object-oriented programming yet, here's some scaffolding. Treat this as a preview for when we get to OOP, but don't stress about the details yet. You need to complete the functions and test them, though.)

<!-- grade_id: sphere, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import math
pi = math.pi
# think of classes as templates or prototypes for specific objects
class Sphere:
    """ The Sphere object has a radius and methods for computing things about it.
    
    Args: (the arguments/parameters passed to the __init__ function)
      radius (float): The radius of the desired sphere in units; this is used for 
        calculating its volume and surface area.
        
    Attributes: (the bits of data that we store about the Sphere)
      radius (float): The radius of the sphere. (This line isn't very interesting, I know.)
    
    """
    def __init__(self, radius):
        # __init__ creates an "instance" of the Sphere based on a radius you give it
        # You won't call __init__ directly, but it happens automatically when you
        # call Sphere() the same way you would a function.
        #
        # most functions in a class have 'self' as the first parameter, but you
        # mostly ignore that when you use Sphere later.
        #
        # note that it's customary for __init__ to not have a docstring, but if it
        # does anything unusual, be sure to add some comments
        self.radius = radius
        
        

    def volume(self):
        """ Returns the volume of the sphere based on its radius.
        
        Args: we don't think about `self` as an argument, so we can omit Args in this case
            
        Returns:
          float: volume of the sphere in units**3
        
        """
        # note: if you want to refer to the radius of the sphere, use self.radius
        volume = (4 * pi * self.radius**3)
        
        return volume

        # return ...  # Note that I want you to return the value, not print it
        
    def surface_area(self):
        """ Returns the surface area of a sphere based on its radius.
        Args: None
        Returns:
        float: surface area of the sphere in units **2
        
        """
        surface_area = (4*pi*self.radius**2)
        return surface_area

```

<!-- grade_id: sphere_tests, points: 2, grade: true, solution: false -->
```python
my_sphere = Sphere(3)  # creates a specific sphere with radius=3

# creates a much larger sphere; these are two separate "instances"
# of Sphere objects
my_other_sphere = Sphere(27)  

assert abs(my_sphere.volume() - 339.292006) < 1e-4  # since comparing floating point numbers, avoid checking exact equality
assert abs(my_sphere.surface_area() - 37.699111) < 1e-4

assert Sphere(0.5).radius == 1

```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-7-8599e91b1f2f> in <module>
      6 
      7 assert abs(my_sphere.volume() - 339.292006) < 1e-4  # since comparing floating point numbers, avoid checking exact equality
----> 8 assert abs(my_sphere.surface_area() - 37.699111) < 1e-4
      9 
     10 assert Sphere(0.5).radius == 1

AssertionError:
```

```python
my_sphere.volume()
my_sphere.surface_area()
Sphere(0.5).radius
```

Output:
```
0.5
```

<a id='random'></a>
## Pseudo-random numbers: `random`

You should know all about pseudo-random numbers and variates from OPER 561 Discrete-Event Simulation. Using randomization in Python is similar to Simio.

```python
import random

# a random float from 0 to 1 (excluding 1)
random.random()

pets = ["cat", "dog", "fish"]

# a random element from a sequence
random.choice(pets)

# shuffle a list (in place--`pets` gets changed)
random.shuffle(pets)

# note its possible `pets` will get shuffled into the original 
# order, so run this cell again if that happens
print(pets)

# a random integer from 1 to 10 (inclusive--unusual for Python)
random.randint(1, 10)
```

Output:
```
['fish', 'cat', 'dog']
7
```

You can use the `dir` function on `random` to see your other options for random number generation.

When we load the `random` module, we can seed it before we start generating values. We can think of this as picking a place in the pseudo-random sequence where we want to start. We normally want to start in a different place every time – by default, the module is seeded with a value taken from the system clock. If we want to reproduce the same random sequence multiple times – for example, inside a unit test – we need to pass the same integer or string as parameter to `seed` each time:

```python
# set a predictable seed
random.seed(3)
print(random.random())
print(random.random())
print(random.random())
```

Output:
```
0.23796462709189137
0.5442292252959519
0.36995516654807925
```

```python
# now try it again
random.seed(3)
print(random.random())
print(random.random())
print(random.random())
```

Output:
```
0.23796462709189137
0.5442292252959519
0.36995516654807925
```

```python
# and now try a different seed
random.seed("something completely different")
print(random.random())
print(random.random())
print(random.random())
```

Output:
```
0.5957576038090177
0.4635135210895731
0.8749953767500402
```

Note that the seed is not exactly the same thing as the random number streams you worked with in Simio. Python's `random` module doesn't really support streams by default, but you can improvise or use a different module for randomness; if you work requires this level of control, ask for more information.

Also, once we learn about NumPy, we'll tend to favor its own random number capabilities instead of the built-in `random` module. But `random` is swell for simple programs.

## <font color="blue">Exercise 3</font>

Write a program which randomly picks an integer from 1 to 100. Your program should prompt the user for guesses – if the user guesses incorrectly, it should print whether the guess is too high or too low. If the user guesses correctly, the program should print how many guesses the user took to guess the right answer. You can assume that the user will enter valid input.

We can use the `input()` function to get user input; run `help(input)` for details. This function works at both the command line and in notebooks. Here's some scaffolding.

<!-- grade_id: secret_number, points: 0, grade: false, solution: true -->
```python
import random
import math
random.seed()

secret_number = random.randint(1,100)
guess = None  # this initializes the guess from the user

# YOUR CODE HERE

while guess != secret_number:
    # input() shows a message box/command prompt for user input and returns
    # whatever they type before hitting enter. It will always return a string.
    guess = int(input("Guess a number from 1 to 100: "))
    
    
    
    if guess == secret_number:
    
        print("That's Correct!")
        
        # we could rewrite this loop to make the `break` completely unneccessary...
        # once you solve this using `break`, why not give that a try?
        break  

    if guess < secret_number:
        print("Too low!")
    else:
        print("Too high!")
```

Output:
```
Guess a number from 1 to 100: 7
Too low!
Guess a number from 1 to 100: 10
Too low!
Guess a number from 1 to 100: 50
Too low!
Guess a number from 1 to 100: 75
Too low!
Guess a number from 1 to 100: 87
Too high!
Guess a number from 1 to 100: 80
Too low!
Guess a number from 1 to 100: 84
Too low!
Guess a number from 1 to 100: 85
That's Correct!
```

<!-- grade_id: cell-f2959111e1536568, points: 2, grade: true, solution: false -->
```python
# There are no automated tests for this one. We're on our own.
```

<a id='re'></a>
## Matching string patterns: `re`

The `re` module allows us to write regular expressions. Regular expressions are a mini-language for matching strings, and can be used to find and possibly replace text. If you learn how to use regular expressions in Python, you will find that they are quite similar to use in other languages. (Regular expressions exist separately from Python. You can even use some for sophisticated find-replace in Word!)

The full range of capabilities of regular expressions is quite extensive, and they are often criticised for their potential complexity, but with the knowledge of only a few basic concepts we can perform some very powerful string manipulation easily.

> Regular expressions are good for use on plain text, but a bad fit for parsing more structured text formats like XML – you should always use a more specialised parsing library for those.

The Python documentation for the `re` module not only explains how to use the module, but also contains a reference for the complete regular expression syntax which Python supports.

### A regular expression primer

A regular expression is a string which describes a pattern. This pattern is compared to other strings, which may or may not match it. A regular expression can contain normal characters (which are treated literally as specific letters, numbers or other symbols) as well as special symbols which have different meanings within the expression.

Because many special symbols use the backslash (`\`) character, we often use raw strings to represent regular expressions in Python. This eliminates the need to use extra backslashes to escape backslashes, which would make complicated regular expressions much more difficult to read. If a regular expression doesn’t contain any backslashes, it doesn’t matter whether we use a raw string or a normal string.

Here are some very simple examples:

```python
# this regular expression contains no special symbols
# it won't match anything except 'cat'
"cat"

# a . stands for any single character (except the newline, by default)
# this will match 'cat', 'cbt', 'c3t', 'c!t' ...
"c.t"

# a * repeats the previous character 0 or more times
# it can be used after a normal character, or a special symbol like .
# this will match 'ct', 'cat', 'caat', 'caaaaaaaaat' ...
"ca*t"
# this will match 'sc', 'sac', 'sic', 'supercalifragilistic' ...
"s.*c"

# + is like *, but the character must occur at least once
# there must be at least one 'a'
"ca+t"

# more generally, we can use curly brackets {} to specify any number of repeats
# or a minimum and maximum
# this will match any five-letter word which starts with 'c' and ends with 't'
"c.{3}t"
# this will match any five-, six-, or seven-letter word ...
"c.{3,5}t"

# One of the uses for ? is matching the previous character zero or one times
# this will match 'http' or 'https'
"https?"

# square brackets [] define a set of allowed values for a character
# they can contain normal characters, or ranges
# if ^ is the first character in the brackets, it *negates* the contents
# the character between 'c' and 't' must be a vowel
"c[aeiou]t"
# this matches any character that *isn't* a vowel, three times
"[^aeiou]{3}"
# This matches an uppercase UCT student number
"[B-DF-HJ-NP-TV-Z]{3}[A-Z]{3}[0-9]{3}"

# we use \ to escape any special regular expression character
# this would match 'c*t'
r"c\*t"
# note that we have used a raw string, so that we can write a literal backslash

# there are also some shorthand symbols for certain allowed subsets of characters:
# \d matches any digit
# \s matches any whitespace character, like space, tab or newline
# \w matches alphanumeric characters -- letters, digits or the underscore
# \D, \S and \W are the opposites of \d, \s and \w

# we can use round brackets () to *capture* portions of the pattern
# this is useful if we want to search and replace
# we can retrieve the contents of the capture in the replace step
# this will capture whatever would be matched by .*
"c(.*)t"

# ^ and $ denote the beginning or end of a string
# this will match a string which starts with 'c' and ends in 't'
"^c.*t$"

# | means "or" -- it lets us choose between multiple options.
"cat|dog"
```

Output:
```
'cat|dog'
```

### Using the `re` module

Now that we have seen how to construct regular expression strings, we can start using them. The `re` module provides us with several functions which allow us to use regular expressions in different ways:

* `search` searches for the regular expression inside a string – the regular expression will match if any subset of the string matches.
* `match` matches a regular expression against the entire string – the regular expression will only match if the whole string matches. `re.match('something', some_string)` is equivalent to `re.search('^something$', some_string)`.
    * In regular expression language, `^` means the start of a string and `$` means the end.
* `sub` searches for the regular expression and replaces it with the provided replacement expression.
* `findall` searches for all matches of the regular expression within the string.
* `split` splits a string using any regular expression as a delimiter.

As you can see, this module provides more powerful versions of some simple string operations: for example, we can also split a string or replace a substring using the built-in `split` and `replace` methods – but we can only use them with fixed delimiters or search patterns and replacements. With `re.sub` and `re.split` we can specify variable patterns instead of fixed strings.

All of the functions take a regular expression as the first parameter. `match`, `search`, `findall`, and `split` also take the string to be searched as the second parameter – but in the `sub` function this is the third parameter, the second being the replacement string. All the functions also take an keyword parameter which specifies optional flags, which we will discuss shortly.

`match` and `search` both return match objects which store information such as the contents of captured groups. `sub` returns a modified copy of the original string. `findall` and `split` return a list of strings.

Here are some usage examples:

```python
import re

# match and search are quite similar
print(1, re.match("c.*t", "cravat")) # this will match
print(2, re.match("c.*t", "I have a cravat")) # this won't
print(3, re.search("c.*t", "I have a cravat")) # this will
```

Output:
```
1 <re.Match object; span=(0, 6), match='cravat'>
2 None
3 <re.Match object; span=(9, 15), match='cravat'>
```

```python
# We can use a static string as a replacement...
print(re.sub("lamb", "squirrel", "Mary had a little lamb."))

# Or we can capture groups, and substitute their contents back in.
print(re.sub("(.*) (BITES) (.*)", r"\3 \2 \1", "DOG BITES MAN"))
```

Output:
```
Mary had a little squirrel.
MAN BITES DOG
```

```python
# count is a keyword parameter which we can use to limit replacements
print(re.sub("a", "b", "aaaaaaaaaa"))
print(re.sub("a", "b", "aaaaaaaaaa", count=1))
```

Output:
```
bbbbbbbbbb
baaaaaaaaa
```

```python
# Here's a closer look at a match object.
my_match = re.match("(.*) (BITES) (.*)", "DOG BITES MAN")
print(my_match.groups())
print(my_match.group(1))
```

Output:
```
('DOG', 'BITES', 'MAN')
DOG
```

```python
# note that the first match in the group is 1, not 0...
# this is because group(0) returns the entire match:
print(my_match.group(0))
```

Output:
```
DOG BITES MAN
```

```python
# We can name groups.
my_match = re.match("(?P<subject>.*) (?P<verb>BITES) (?P<object>.*)", "DOG BITES MAN")
print(my_match.group("subject"))
print(my_match.groupdict())

# We can still access named groups by their positions.
print(my_match.group(1))
```

Output:
```
DOG
{'subject': 'DOG', 'verb': 'BITES', 'object': 'MAN'}
DOG
```

```python
# Sometimes we want to find all the matches in a string.
print(re.findall("[^ ]+@[^ ]+", "Bob <bob@example.com>, Jane <jane.doe@example.com>"))

# Sometimes we want to split a string.
print(re.split(", *", "one,two,  three, four"))
```

Output:
```
['<bob@example.com>,', '<jane.doe@example.com>']
['one', 'two', 'three', 'four']
```

### Greed

Regular expressions are greedy by default – this means that if a part of a regular expression can match a variable number of characters, it will always try to match as many characters as possible. That means that we sometimes need to take special care to make sure that a regular expression doesn’t match too much. For example:

```python
# this is going to match everything between the first and last double quote "
# (i.e., the whole string)
# but that's not what we want!
print(re.findall('".*"', '"one" "two" "three" "four"'))

# This is a common trick
print(re.findall('"[^"]*"', '"one" "two" "three" "four"'))

# We can also use ? after * or other expressions to make them *not greedy*
print(re.findall('".*?"', '"one" "two" "three" "four"'))
```

Output:
```
['"one" "two" "three" "four"']
['"one"', '"two"', '"three"', '"four"']
['"one"', '"two"', '"three"', '"four"']
```

**Matching other special chararacters:**

You often need to "escape" special characters (most non-alphanumeric characters, but especially parens and brackets) by preceeding them with a backslash:

```python
my_string = "Hello (hi)."
print(re.search("(hi)", my_string))  # while this finds 'hi', it misses the parens
print(re.search("\(hi\)", my_string)) # this gets the parens
print(re.search(r"(hi)", my_string))  # this doesn't and I'm not sure why
```

Output:
```
<re.Match object; span=(7, 9), match='hi'>
<re.Match object; span=(6, 10), match='(hi)'>
<re.Match object; span=(7, 9), match='hi'>
```

## <font color="blue">Exercise 4</font>

*Flipside ver. 2* - Write a function which takes a string which contains two words separated by any non-zero amount and type of whitespace---assume spaces & tabs--and returns a string in which the words are swapped around and the whitespace is preserved. Use the `re` module. The character class `[\t ]` is a good way to match whitespace (`\t` is the tab character).

<!-- grade_id: flipside2, points: 0, grade: false, solution: true -->
```python
import re
def flipside2(s):
    """...
    
    """
    re
```

<!-- grade_id: flipside2_tests, points: 2, grade: true, solution: false -->
```python
assert flipside2('nice code') == 'code nice'
assert flipside2('Air      Power') == 'Power      Air'
assert flipside2('hidden  	  tab') == 'tab  	  hidden'
```

<a id='csv'></a>
## Parsing CSV files: `csv`

CSV stands for *comma-separated values* – it’s a very simple file format for storing tabular data. Most spreadsheets can easily be converted to and from CSV format.

In a typical CSV file, each line represents a row of values in the table, with the columns separated by commas. Field values are often enclosed in double quotes, so that any literal commas or newlines inside them can be escaped:

```python
"one","two","three"
"four, five","six","seven"
```

Output:
```
('four, five', 'six', 'seven')
```

Python’s `csv` module takes care of all this in the background, and allows us to manipulate the data in a CSV file in a simple way, using the `reader` and `writer` classes:

```python
import csv

# this will create pets.csv in the current directory
with open('pets.csv', 'w', newline='') as f:
    pets_writer = csv.writer(f)
    pets_writer.writerow(['Fluffy', 'cat'])
    pets_writer.writerow(['Max', 'dog'])

print('done writing')

with open("pets.csv") as f:
    pets_reader = csv.reader(f)
    for row in pets_reader:
        print(row)
```

Output:
```
done writing
['Fluffy', 'cat']
['Max', 'dog']
```

(A note on the `newline=''` argument: Windows and non-Windows operating systems handle line endings a little differently. Without the `newline=''` parameter, you'll often get extra blank lines appearing in the output. Try removing that parameter and re-running. For more details, you can read [this](https://docs.python.org/3.6/library/csv.html#id3).)

There is no single CSV standard – the comma may be replaced with a different delimiter (such as a tab or a pipe), and a different quote character may be used. Both of these can be specified as optional keyword parameters to `reader`.

## <font color="blue">Exercise 5</font>

Open a CSV file which contains three columns of integers; you'll need to create one for your testing, but don't upload it as part of your assignment. 

Write out the data to a *new* CSV file, swapping around the second and third columns and adding a fourth column which contains the numerical sum of the first three. (So if your numbers get read in as strings, you'll need to convert them to integers.)

Create a `list` of (row number, row sum) tuples; row sum is what you wrote to the 4th column of the CSV file. Consider using a [list comprehension](https://www.pythonlikeyoumeanit.com/Module2_EssentialsOfPython/Generators_and_Comprehensions.html#The-range-generator) and/or the [zip function](https://www.pythonlikeyoumeanit.com/Module2_EssentialsOfPython/Itertools.html#Python%E2%80%99s-%E2%80%9CItertools%E2%80%9D).

<!-- grade_id: edit_csv, points: 0, grade: false, solution: true -->
```python
import csv
import pandas as pd

def edit_csv(input_filename, output_filename):
    """
    
    """    
        df = pd.Dataframe
            
                
        
        
```

```python
import pandas as pd
import numpy as np
df = pd.read_csv("Book1.csv", 
                  names=["col1", "col2", "col3"])
print(df)
df = df[['col1','col3','col2']]
print(df)
df['Row Sum'] = df.apply(np.sum, axis=1)
print(df)

df['index1'] = df.index

new_list
#new_list = list(index, df['Row Sum'])

#df = df.to_csv('Book2.csv')


```

Output:
```
col1  col2  col3
0     42    68     4
1     49    22    75
2     65    95    82
3     96    62    82
4     93    92    52
5     80    88    11
6     80    60    87
7     65    99    31
8     39    55    11
9     51    32    20
10    98    88    49
11    33    16    61
12    84    94    40
13    37    13    87
14    48    33    45
15    10    91    20
16    89    66    54
17    52    68     2
18    21    94    86
19    51    53    14
20    96    72    19
21    30    11    58
22    70    76    10
23    37    33    82
24    81    34    19
25    47    31    58
26    95     4    52
27    24    99    85
28   100    17    29
29    20     8    23
30    54    10    33
31    75    41    32
32    15    42    33
33     9    22    48
34    40     5    52
35    28    30    33
36    33    26    64
37    65    47    80
38    98    76    64
39    60    97     2
    col1  col3  col2
0     42     4    68
1     49    75    22
2     65    82    95
3     96    82    62
4     93    52    92
5     80    11   
```

<!-- grade_id: edit_csv_tests, points: 2, grade: true, solution: false -->
```python
#
# create an input file for testing
#
with open('_testinput.csv', 'w', newline='') as f:
    f.write("""1,2,3
4,5,6""")  # using """ lets a string be spread across multiple lines and include newline characters

#
# test the function by manually working with csv data
#
edit_csv('_testinput.csv', '_testoutput.csv')

with open('_testoutput.csv', 'r') as f:
    res = f.readlines()
    
# readlines returns a list, which I want to join into one string; 
# strip removes extra whitespace
res = ''.join(res).strip()
    
desired_output = """1,3,2,6
4,6,5,15"""

assert res == desired_output

#
# cleanup after testing
#
import os
try:
    os.remove('_testinput.csv')
except:
    pass  # if we can't remove it, oh well
try:
    os.remove('_testoutput.csv')
except:
    pass
```

<a id='questions'></a>
# Questions

1) Briefly define the decision-making phases of intelligence, design, choice, and implementation.

<!-- grade_id: question1, points: 2, grade: true, solution: true -->


2) There is a trade-off between model accuracy and cost. Come up with an example of this for some real-world context.

<!-- grade_id: question2, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

3) In Turban Fig. 2.1, a diagram of the system & its environment, I've highlighted "Decision maker". What do you think is their impact on designing a DSS? How might designing a DSS differ from writing programs for your own problem solving?

<!-- grade_id: question3, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

4) We tend to overemphasize decision support in the Choice Phase, because it leads to choosing a solution. But the Turban reading discusses the other phases, too. Briefly describe a decision context (of your own design or from elsewhere) where support in the Intelligence Phase could be useful. And also for the Design Phase.

<!-- grade_id: question4, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

5) In the csv example, we saw

`with open("pets.csv") as f:
    pets_reader = csv.reader(f)`

This is an example of managing a resource (here, a file) using a [context manager](https://book.pythontips.com/en/latest/context_managers.html). According to that reference, what's the most important benefit of using a context manager for working with a file? What would be a functionally equivalent, but worse, way of using pets.csv?

<!-- grade_id: question5, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

```python
# code answer if you need it
```

# Completion

You should save and submit this notebook file to Canvas under the [Lab 2 Assignment](https://lms.au.af.edu/courses/24019/assignments/202806) prior to leaving class.

Assessment:

* (10 pts) Each of the five exercises should:
    * have an appropriate docstring
    * have been tested at least once
    * run without error for expected inputs
    * produce correct output
* (10 pts) Each of the Questions are answered appropriately.
* If you worked with others, list their names below.